---
title: Getting Started
description: Creating Computer Programs
---
## What is Programming?

```yaml
type: PureMultipleChoiceExercise
key: e644f1bcc3
```
**Computer programming** is a way for people to provide precise instructions to computers. In business, people create standard processes and procedures to ensure work is performed efficiently, and likewise, people can create computer programs so computers can complete tasks effectively.

![Computer Programming Meme](https://i.imgur.com/kxkJCeN.png])<br/><br/>
[View on imgur](https://imgur.com/gallery/kxkJCeN)

By viewing this course, you are benefitting from computer programs other people have created. For example, there are computer programs ensuring that each pixel in your computer screen is displaying the correct color to produce the image on the screen. When you move your mouse or trackpad, you see the cursor move on the screen because your computer contains computer programs that instruct it to reach to your mouse or trackpad movements. A key to identifying a computer program is to ask: "What/who is executing the instructions?" Computer programs are executed be computers or devices rather than human.

Which of the following is __NOT__ a computer program?

`@possible_answers`
- Microsoft Excel
- [Instructions posted on the company intranet for uploading approval forms (correct one)]
- "Auto-Defrost" option on the microwave
- [Device driver](https://en.wikipedia.org/wiki/Device_driver) for external microphone

`@feedbacks`
- That is a computer program. Software applications are computer programs.
- That is correct. Those instructions are performed by people who are uploading the approval forms.
- That is a computer program. The "Auto-Defrost" option contains instructions that are executed by the microwave when you select you.
- That is a computer program. Device drivers contain instruction that tell your computer how to work with the connected device.

---
## Python

```yaml
type: PureMultipleChoiceExercise
key: 5acfc43ca9
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
